##### Q5
จงเขียนโปรแกรมเพื่อรับตัวเลขหนึ่งตัวระหว่าง [0,9] จากนั้นแสดงตัวเลขตั้งแต่ 1-100 ซึ่งมีเลขที่รับเข้ามาปรากฏอยู่ <br>

สิ่งที่ต้องทำคือ
- สร้างฟังก์ชัน create_list(number) **(รับค่า number ที่เป็น integer)** เพื่อสร้างลิสต์ที่มีเลขที่รับเข้ามาปรากฏอยู่ คืนค่าเป็นลิสต์

### ตัวอย่างการทำงาน:
Enter a nember between [0,9] : **0** <br>
[10, 20, 30, 40, 50, 60, 70, 80, 90, 100] <br>

### ตัวอย่างการทำงาน:
Enter a number between [0,9]: **3** <br>
[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93] <br>

### ตัวอย่างการทำงาน:
Enter a number between [0,9]: **4** <br>
[4, 14, 24, 34, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 54, 64, 74, 84, 94] <br>

### ตัวอย่างการทำงาน:
Enter a number between [0,9]: **8** <br>
[8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98] <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def create_list(number):
    lst = [i for i in range(1,101) if str(number) in str(i)]
    return lst

def run():
    num = input('Enter a number between [0,9]:')
    print(create_list(num))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter a number between [0,9]: 4
[4, 14, 24, 34, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 54, 64, 74, 84, 94]
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output):
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
def test_easy_print(inputs,expected_output):
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert expected_output in mock_stdout.getvalue(), mock_stdout.getvalue()
```

<!-- grade_id: easy_print, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
test_easy_print(inputs=['2'],expected_output='2')
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-0c38a636281d4f0d, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['0'],"[10, 20, 30, 40, 50, 60, 70, 80, 90, 100]\n")
test(['4'],"[4, 14, 24, 34, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 54, 64, 74, 84, 94]\n")
test(['8'],"[8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98]\n")
test(['1'],"[1, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 21, 31, 41, 51, 61, 71, 81, 91, 100]\n")
test(['3'],"[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93]\n")
test(['9'],"[9, 19, 29, 39, 49, 59, 69, 79, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99]\n")
print('ok')
```

Output:
```
ok
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน create_list'''
assert create_list(5) == [5, 15, 25, 35, 45, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 65, 75, 85, 95]
assert create_list(6) == [6, 16, 26, 36, 46, 56, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 76, 86, 96]
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-59f35ec9e84d2ff2, points: 1, grade: true, solution: false -->
```python
'''เทสลับ1'''
```

Output:
```
'เทสลับ1'
```

<!-- grade_id: cell-183d686f87a66cee, points: 1, grade: true, solution: false -->
```python
'''เทสลับ2'''
```

Output:
```
'เทสลับ2'
```

```python

```