## Q4
จงเขียนฟังก์ชัน sum_digits ซึ่งรับ string ที่ประกอบไปด้วยตัวอักษรและตัวเลข แล้วคืนค่าผลบวกของตัวเลขทั้งหมดที่ปรากฏอยู่ใน string นั้น

### ตัวอย่างการทำงาน: 

Enter your string: zz1xa2d3 <br>
6

### ตัวอย่างการทำงาน: 

Enter your string: 5432d <br>
14

### ตัวอย่างการทำงาน: 

Enter your string: Hello <br>
0

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def sum_digits(text) :    
    num = [int(i) for i in text if i.isdigit() == True]
    return sum(num)
    
def run():
    str = input('Enter your string: ')
    print(sum_digits(str))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter your string:  5432s
14
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
def test_easy_print(inputs,expected_output):
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert expected_output in mock_stdout.getvalue(), mock_stdout.getvalue()
```

<!-- grade_id: cell-9020e4269dec96b8, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
test_easy_print(inputs=['1'],expected_output='1')
print('ok')
```

Output:
```
ok
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Example'''
test(['zz1xa2d3'],'6\n') 
test(['5432d'],'14\n') 
test(['Hello'],'0\n')
test(['Compog2301172'],'16\n')
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-208f643fceb644e5, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน sum_digits'''
assert sum_digits('5G1xa2d3') == 11
assert sum_digits('8c1o') == 9
assert sum_digits('Hel9lo8Worl8dlnw') == 25
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-0c38a636281d4f0d, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```

Output:
```
'เทสลับ'
```

```python

```